---
title       : Regression
description : UNDER CONSTRUCTION

--- type:NormalExercise lang:r xp:100 skills:1 key:0e8ad24c8e

## Exercise 1

Under construction.

*** =instructions
- instruction 1
- instruction 2

*** =hint
- hint 1
- hint 2


*** =pre_exercise_code
```{r}
# pre exercise code here
```

*** =sample_code
```{r}
# sample code here
```

*** =solution
```{r}
#solution code here
```

*** =sct
```{r}
# submission correctness tests

# example tests:
# test_output_contains("output")
# test_object("object_name")
# test_function("function_name", args=c("arg1"))

# test if the students code produces an error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Good work!")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:0b53d4e57c

## Exercise 2

Under construction.

*** =instructions
- instruction 1
- instruction 2

*** =hint
- hint 1
- hint 2


*** =pre_exercise_code
```{r}
# pre exercise code here
```

*** =sample_code
```{r}
# sample code here
```

*** =solution
```{r}
#solution code here
```

*** =sct
```{r}
# submission correctness tests

# example tests:
# test_output_contains("output")
# test_object("object_name")
# test_function("function_name", args=c("arg1"))

# test if the students code produces an error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Good work!")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:f59dd3081e

## Exercise 3

Under construction.

*** =instructions
- instruction 1
- instruction 2

*** =hint
- hint 1
- hint 2


*** =pre_exercise_code
```{r}
# pre exercise code here
```

*** =sample_code
```{r}
# sample code here
```

*** =solution
```{r}
#solution code here
```

*** =sct
```{r}
# submission correctness tests

# example tests:
# test_output_contains("output")
# test_object("object_name")
# test_function("function_name", args=c("arg1"))

# test if the students code produces an error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Good work!")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:7ec065b7c6

## Exercise 4

Under construction.

*** =instructions
- instruction 1
- instruction 2

*** =hint
- hint 1
- hint 2


*** =pre_exercise_code
```{r}
# pre exercise code here
```

*** =sample_code
```{r}
# sample code here
```

*** =solution
```{r}
#solution code here
```

*** =sct
```{r}
# submission correctness tests

# example tests:
# test_output_contains("output")
# test_object("object_name")
# test_function("function_name", args=c("arg1"))

# test if the students code produces an error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Good work!")

```

--- type:NormalExercise lang:r xp:100 skills:1 key:8d30339288

## Exercise 5

Under construction.

*** =instructions
- instruction 1
- instruction 2

*** =hint
- hint 1
- hint 2


*** =pre_exercise_code
```{r}
# pre exercise code here
```

*** =sample_code
```{r}
# sample code here
```

*** =solution
```{r}
#solution code here
```

*** =sct
```{r}
# submission correctness tests

# example tests:
# test_output_contains("output")
# test_object("object_name")
# test_function("function_name", args=c("arg1"))

# test if the students code produces an error
test_error()

# Final message the student will see upon completing the exercise
success_msg("Good work!")

```

