--- type:NormalExercise lang:r xp:100 skills:1
## Exercise title

Exercise info

*** =instructions
- instruction 1

*** =hint
- hint 1

*** =pre_exercise_code
```{r}
# pre exercise code here
```

*** =sample_code
```{r}
# sample code here
```

*** =solution
```{r}
#solution code here
```

*** =sct
```{r}

# test_output_contains("output")
# test_object("object_name")
# test_function("function_name", args=c("arg1"))

test_error()
success_msg("Good work!")
```
