--- type:NormalExercise lang:r xp:100 skills:1
## Function arguments

During this course you will operate on R's ready-made functions, but for clarity we will look at the definition of a simple function and learn how to use it (*call* it).

Functions usually have one or more arguments, for which the user can define values when calling the function. Multiple arguments are separated with a comma. The arguments have names, which can (and often should) be used when specifying their values.

The recommended style is to specify all but the first argument by name, when calling a function.

*** =instructions
- Look at the definition of the `add_k()` function and execute the function calls that are defined, starting from line 5.
- Create object `my_answer` by directly assigning the output of `add_k()` to it. Call `add_k()` with `number = 10` and `k = 7`.

*** =hint
- The answer is of the form: `my_answer <- add_k(arg1 = ?, arg2 = ?)`

*** =pre_exercise_code
```{r}
add_k <- function(number, k = 5) return(number + k)
```


*** =sample_code
```{r}

# Create a simple function with two arguments, of which 'k' has a default value 
add_k <- function(number, k = 5) return(number + k)

# Call the add_k() function. Can you guess the result?
add_k(number = 5)

# These calls are identical. Try them.
add_k(5, k = -4) # recommended style
add_k(5, -4)     # style usually not recommended
add_k(k = -4, number = 5) # style not recommended

# Call add_k() to create the object 'my_answer'
my_answer <- 


```

*** =solution
```{r}
# Create a simple function with two arguments, of which 'k' has a default value
add_k <- function(number, k = 5) return(number + k)

# Call the add_k() function. Can you guess the result?
add_k(number = 5)

# These calls are identical. Try them.
add_k(5, k = -4) # recommended style
add_k(5, -4)     # style usually not recommended
add_k(k = -4, number = 5) # style not recommended

# Call add_k() to create the object 'my_answer'
my_answer <- add_k(10, k = 7)


```

*** =sct
```{r}
test_object("my_answer")

test_error()
success_msg("Great job!")
```


--- type:NormalExercise lang:r xp:100 skills:1
## Brackets have your data [ under, control ]

It is quite often necessary to access only a part of a dataset. This is called *subsetting* and there are many ways of doing it. The best way will depend on the situation and  preference.  

We have already seen that the columns of a data frame (df) can be accessed via `$`. There is another, more general method: `[ ]`. The brackets are a way of accessing a part of any data structure in R. 

The way to use brackets is by giving the locations of the values you wish to access.


*** =instructions
- Study the examples on how to select parts of a data frame.
- Replace `NULL` so that the object `age_height` contains the third and fourth columns of students2014.
- Modify the code to print out the first five rows of `age_height`.


*** =hint
- Selecting multiple columns is very similar to selecting multiple rows.
- See the first example on how to select multiple rows.


*** =pre_exercise_code
```{r}
students2014 <- read.table("http://www.helsinki.fi/~kvehkala/JYTmooc/JYTOPKYS-data.txt", sep="\t", header=TRUE)
  
  # keep a couple background variables
  students2014 <- students2014[,c("sukup","toita","ika","pituus","kenka","kone")]
  
  # recode kone -variable NA values as factor levels
  students2014$kone <- addNA(students2014$kone)
  
  # choose rows without missing values
  students2014 <- students2014[complete.cases(students2014),]

```

*** =sample_code
```{r}
# students2014 is available

# Choose the rows from 3 to 6 in the data
students2014[3:6, ]

# Look at the value in row 6, column 4
students2014[6, 4]

# Choose the whole column 3 and save it as an object
student_age <- students2014[, 3]

# Choose the third and fourth columns and save them as an object
age_height <- students2014[, NULL]

# Print out the first five rows of age_height
age_height[NULL, ]


```

*** =solution
```{r}
# students2014 is available

# Choose the rows from 3 to 6 in the data
students2014[3:6, ]

# Look at the value in row 6, column 4
students2014[6, 4]

# Choose the whole column 3 and save it as an object
student_age <- students2014[, 3]

# Choose the third and fourth columns and save them as an object
age_height <- students2014[, 3:4]

# Print out the first five rows of age_height
age_height[1:5, ]



```

*** =sct
```{r}
test_object("age_height")

test_error()
success_msg("Good work!")
```

--- type:NormalExercise lang:r xp:100 skills:1
## Vectors (2)

The values in vectors have indices: the first value in a vector has a index `1`, the second value has `2` and so on. 

You can get a single value from a vector by using indices and bracets `[ ]`. Using a single index number inside the bracets gives you a single value from the vector. Using a vector of indices gives you a vector of values from the vector. 

Note that you can also rearrange values in a vector by using indices. Look at the example code to see how indices work. 

*** =instructions
- Create the name vector.
- See how the indices work by executing the example lines.
- Create a new vector `girls` with values `Liisa` and `Anna` from `names` by using indices.
- Create also a new vector `boys` with values `Pekka` and `Matti` from `names` by using indices.

*** =hint
- Note that space between the vector object and bracets produces an error.
- Remember that vectors are created with the function `c()`.
- Indices vectors `c(1,2)` and `c(2,1)` do not produce the same outcome. The order of the values is different.

*** =pre_exercise_code
```{r}
```


*** =sample_code
```{r}
# Let's create a vector
names <- c('Matti', 'Pekka', 'Liisa', 'Anna')

# The first value in the vector
names[1]

# The third and fourth value in the vector
names[c(3, 4)]

# Rearranged
names[c(4, 3)]

# Separate the names vector into two vectors by using indices
girls <-
boys <-

```

*** =solution
```{r}
# Let's create a vector
names <- c('Matti', 'Pekka', 'Liisa', 'Anna')

# The first value in the vector
names[1]

# The third and fourth value in the vector
names[c(3, 4)]

# Rearranged
names[c(4, 3)]

# Separate the names vector into two vectors by using indices
girls <- names[c(3, 4)]
boys <- names[c(2, 1)]

```

*** =sct
```{r}
#error in tests
#ex() %>% check_output_expr("mtcars[4, ]", missing_msg = "Have you used ... ")
test_output_contains("c(0.1, 0.2, 5.84, 0.7)", incorrect_msg = "Did you execute all the lines?")

test_object("names", incorrect_msg = "Did you create `names`?")
test_object("girls", incorrect_msg = "Did you create `girls`?")
test_object("boys", incorrect_msg = "Did you create `boys`?")

test_error()

# Final message the student will see upon completing the exercise
success_msg("Such indices. Wow. Much intelligent.")

```
--- type:NormalExercise lang:r xp:100 skills:1
## Data types and functions

Not all data types do well with certain functions. 

One cannot for example use character vectors when the function needs a numeric one. It is also impossiple to create a vector with multiple data types. 

When using the wrong data type, R produces an error stating what is wrong. 

*** =instructions
- Create `my_vector`.
- Try to calculate the mean of `my_vector` and see what happens. 
- Correct the vector and calculate the mean so that you do not get any errors.

*** =hint
- You can use `str()` to check the data types.
- `chr` is short for character.
- When creating a character vector, one needs to use quotation marks. 

*** =pre_exercise_code
```{r}
```

*** =sample_code
```{r}
# Let's create a vector
my_vector <- c(2, 3, "1", 5)

# Calculate the mean of vector
mean(my_vector)
  
```

*** =solution
```{r}
# Let's create a vector
my_vector <- c(2, 3, 1, 5)

# Calculate the mean of vector
mean(my_vector)

```

*** =sct
```{r}
test_object("my_vector", incorrect_msg = "Did you correct the `my_vector`?")

test_error()

# Final message the student will see upon completing the exercise
success_msg("Creating vectors like a pro. Great work!")

```
